Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Caleb Andree"
COLLABORATORS = ""
```

---

# Lab 7 Text Analysis [Total: 9 points]

The purpose of this assignment is for you to engage with a concrete text analysis task. This will be accomplished through a coding assignment. You will carry out this task in the present notebook, and use the notebook to document the various steps of the exercise and to answer all questions.


## Required skills

This lab will let you practice the following skills:
- Write regular expressions
- Parse date formats
- Parse email addresses
- Make substitutions using regular expressions.

Before you start working on it, review P4E Ch. 11 of the textbook and/or the notebooks posted on ELMS under "Text Analysis". 

A few additional resources can be found here:
- https://docs.python.org/3/library/re.html
- https://www.dataquest.io/blog/regex-cheatsheet/
- http://www.pyregex.com/


## Table of Contents
<ul>
    <li><a href="#Submission-checklist">Submission checklist</a></li>
    <li><a href="#Q1">Question 1</a></li>
    <li><a href="#Q2">Question 2</a></li>
    <li><a href="#Q3">Question 3</a></li>
    <li><a href="#Q4">Question 4</a></li>
</ul>

## Submission checklist

**Points**: 1

Before submitting make sure that:

1. Your name is included above, plus the name of any collaborator you worked with;
2. All Markdown cells you edited are rendering correctly, especially the ones with answers.
3. You have removed any `raise NotImplementedError()` line from your code cells.

## Q1

**Points**: 2

Write a function called `yearofdate` that matches the year from a date timestamp 

Your function should take one parameter, a date string of the format `MM/DD/YYYY`, and it should return the year part (`YYYY`) as an int.

### Examples

For example, if the date string is `01/21/2018`, then your function should return `2018` (as an int).

---

This question is worth 2 points:
- 1 point for returning the correct year (automatically graded),
- 1 point for making use of regular expressions (graded manually by us).

<!-- grade_id: cell-f34beeb61e2454e4, points: 1, grade: true, solution: true -->
```python
import re

# YOUR CODE HERE
def yearofdate(date):
    pattern = r'\d{4}'
    year_matches = re.findall(pattern, date)
    return int(year_matches[0])
#raise NotImplementedError()
```

Use the cell below to run your function and see what it returns. You may want to test different dates to see if your code returns the correct answer.

```python
pattern = '12/12/2012'
yearofdate(pattern)
```

Output:
```
2012
```

<!-- grade_id: cell-e19a284c643d6c2f, points: 1, grade: true, solution: false -->
```python
# autograder goes here
import pandas as pd

# Create test cases
test_cases = ['01/21/2018', '12/12/2012', '03/03/2018']
expected_results = [2018, 2012, 2018]

# Autograder
for i, date in enumerate(test_cases):
    student_answer = yearofdate(date)
    correct_answer = expected_results[i]
    try:
        assert student_answer == correct_answer
        print(f"Test {i + 1}: Passed")
    except AssertionError:
        print(f"Test {i + 1}: Failed")
        print(f"Student Answer: {student_answer}")
        print(f"Expected Answer: {correct_answer}")

# Print a summary
print("All tests passed! 👍")
```

Output:
```
Test 1: Passed
Test 2: Passed
Test 3: Passed
All tests passed! 👍
```

## Q2

**Points**: 2

Write a function called `monthofdate` that matches the month from a date timestamp 

Your function should take one parameter, a date string of the format `<MONTH> <DAY>, YYYY`, where `<MONTH>` is the name of a month, `<DAY>` is the day number, and `YYYY` is the year, and it should return the month as an int (`1` = January, etc.).

### Examples

For example, if the date string is `March 8, 2017`, then your function should return `3` (as an int).

---

This question is worth 2 points:
- 1 point for returning the correct month (automatically graded),
- 1 point for making use of regular expressions (graded manually by us).

<!-- grade_id: cell-ceaca6c1f47726bc, points: 1, grade: true, solution: true -->
```python
import re
from datetime import datetime

# YOUR CODE HERE
def monthofdate(date):
    pattern = r'(\w+) \d+, \d{4}'
    match = re.search(pattern, date)
    month_name = match.group(1)

    
    month_number = datetime.strptime(month_name, '%B').month
    return month_number


#raise NotImplementedError()
```

Use the cell below to run your function and see what it returns. You may want to test different dates to see if your code returns the correct answer.

```python
pattern = 'March 8, 2017'
monthofdate(pattern)
```

Output:
```
3
```

<!-- grade_id: cell-29bc3f2ab4cc8d86, points: 1, grade: true, solution: false -->
```python
import pandas as pd

# Create test cases
test_cases = ['March 8, 2017', 'January 15, 2018', 'May 3, 2017']
expected_results = [3, 1, 5]

# Autograder
for i, date in enumerate(test_cases):
    student_answer = monthofdate(date)
    correct_answer = expected_results[i]

    try:
        assert student_answer == correct_answer
        print(f"Test {i + 1}: Passed")
    except AssertionError:
        print(f"Test {i + 1}: Failed")
        print(f"Student Answer: {student_answer}")
        print(f"Expected Answer: {correct_answer}")

# Print a summary
print("All tests passed! 👍")
```

Output:
```
Test 1: Passed
Test 2: Passed
Test 3: Passed
All tests passed! 👍
```

## Q3

**Points**: 2

Write a function called `emailuppercase` that converts the domain name of an email address to uppercase characters. 

Your function should take one parameter, an email address and should return the same address, but with the domain name part (i.e., what is after the `@`) in upper case.

### Examples

For example, if the email is `gciampag@umd.edu`, then your function should return `gciampag@UMD.EDU`.

### Hint

Consider using the `re.sub` function, see documentation: https://docs.python.org/3.11/library/re.html#re.sub

---

This question is worth 2 points:
- 1 point for returning the email address (automatically graded),
- 1 point for making use of regular expressions (graded manually by us).

<!-- grade_id: cell-ff1732a5bd404f06, points: 1, grade: true, solution: true -->
```python
import re
# YOUR CODE HERE
def emailuppercase(email):
    pattern = r'@([a-zA-Z0-9.-]+)$'
    modified_email = re.sub(pattern, lambda match: "@" + match.group(1).upper(), email)
    return modified_email
#raise NotImplementedError()
```

Use the cell below to run your function and see what it returns. You may want to test different dates to see if your code returns the correct answer.

```python
pattern = 'gciampag@umd.edu'
emailuppercase(pattern)
```

Output:
```
'gciampag@UMD.EDU'
```

<!-- grade_id: cell-52e0100f41899e6f, points: 1, grade: true, solution: false -->
```python
# autograder
import re

# Create test cases
test_cases = ['gciampag@umd.edu', 'student1@terpmail.umd.edu', 'dr.evil@ox.ac.uk']
expected_results = ['gciampag@UMD.EDU', 'student1@TERPMAIL.UMD.EDU', 'dr.evil@OX.AC.UK']

# Autograder
for i, email in enumerate(test_cases):
    student_answer = emailuppercase(email)
    correct_answer = expected_results[i]

    try:
        assert student_answer == correct_answer
        print(f"Test {i + 1}: Passed")
    except AssertionError:
        print(f"Test {i + 1}: Failed")
        print(f"Student Answer: {student_answer}")
        print(f"Expected Answer: {correct_answer}")

print("All tests passed! 👍")
```

Output:
```
Test 1: Passed
Test 2: Passed
Test 3: Passed
All tests passed! 👍
```

## Q4

**Points**: 2

Write a function called `ssnredacter` that redacts a social security number. 

Your function should take one parameter, a string with a social security number with format `NNN-NN-NNNN` and should return the same number, but with the first 5 digits replaced with asterisks (`*`).

### Examples

For example, if the SSN is `111-11-1111`, then your function should return `***-**-1111`.

### Hint

Consider using the `re.sub` function, see documentation: https://docs.python.org/3.11/library/re.html#re.sub

---

This question is worth 2 points:
- 1 point for returning the correct SSN (automatically graded),
- 1 point for making use of regular expressions (graded manually by us).

<!-- grade_id: cell-5f64551f822c8d2d, points: 1, grade: true, solution: true -->
```python
import re

# YOUR CODE HERE
def ssnredacter(ssn):
    ssn_pattern = r'^(\d{3}-\d{2}-)(\d{4})$'
    redacted_ssn = re.sub(ssn_pattern, r'***-**-\2', ssn)
    return redacted_ssn

#raise NotImplementedError()
```

Use the cell below to run your function and see what it returns. You may want to test different dates to see if your code returns the correct answer.

```python
pattern = '111-11-1111'
ssnredacter(pattern)
```

Output:
```
'***-**-1111'
```

<!-- grade_id: cell-89957d4384da9979, points: 1, grade: true, solution: false -->
```python
import re

# Create test cases
test_cases = ['111-11-1111', '999-19-1919', '888-12-3434']
expected_results = ['***-**-1111', '***-**-1919', '***-**-3434']

# Autograder
for i, ssn in enumerate(test_cases):
    student_answer = ssnredacter(ssn)
    correct_answer = expected_results[i]

    try:
        assert student_answer == correct_answer
        print(f"Test {i + 1}: Passed")
    except AssertionError:
        print(f"Test {i + 1}: Failed")
        print(f"Student Answer: {student_answer}")
        print(f"Expected Answer: {correct_answer}")

print("All tests passed! 👍")
```

Output:
```
Test 1: Passed
Test 2: Passed
Test 3: Passed
All tests passed! 👍
```